Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Yanjun Zhang"

```

---

# Using classes in Python

* Date: 2023-02-20
* Deadline: 2023-02-27

## 1. Create class for 3D vector 

Create a class for storing a 3D vector (x, y, z coordinates), and which provides a set of functions for manipulating the coordinates.  

* Implement class Vector3D which has the following properties:

### a. instance attributes

- coord_x stores x coordinate
- coord_y stores y coordinate
- coord_z stores z coordinate
    
### b. instance methods

- length() : computes length of vector* 
- renormalize()   : renormalizes vector to unit length vector*
- for square root use sqrt function from math library

### c. Overloaded operations  

- addition operator(+)     : implement `__add__` method 
- substraction operator(-) : implement `__sub__` method 

    

<!-- grade_id: cell-106104681ce80401, points: 0, grade: false, solution: true -->
```python
import math

class Vector3D:
    """
    Class for manipulation of 3D-vectors
    
    instance attributes:
        coord_x, coord_y, coord_z (float)
    
    class methods
        length (float): returns Euclidean norm of vector
        normalize (None): scales coordinates so that resulting vector has length 1
        __add__ (Vector3D) returns new vector as vector sum of two vectors
        __sub__ (Vector3D) returns new vector as vector difference of two vectors
        __repr__: string representation of vector

    """
    # YOUR CODE HERE
    ################
    def __init__(self,coord_x,coord_y,coord_z):
        self.coord_x = float(coord_x)
        self.coord_y = float(coord_y)
        self.coord_z = float(coord_z)

    def length(self):
        return math.sqrt(self.coord_x**2+self.coord_y**2+self.coord_z**2)

    def renormalize(self):
        length = self.length()
        if length == 0:
            raise ValueErro("cannot normalize a zero-length vector")
        self.coord_x /= length
        self.coord_y /= length
        self.coord_z /= length

    def __add__(self,oth):
        return Vector3D(
            self.coord_x + oth.coord_x,
            self.coord_y + oth.coord_y,
            self.coord_z + oth.coord_z,
        )

    def __sub__(self,oth):
        return Vector3D(
            self.coord_x - oth.coord_x,
            self.coord_y - oth.coord_y,
            self.coord_z - oth.coord_z,
        )

    def __repr__(self):
        return f"Vector3D({self.coord_x},{self.coord_y},{self.coord_z})"
    

    
    
```

### a. instance attributes

* coord_x stores x coordinate
* coord_y stores y coordinate
* coord_z stores z coordinate


<!-- grade_id: cell-e8becee65f92a575, points: 1, grade: true, solution: false -->
```python
# Verify Vector3D constructor
from pytest import approx
vector = Vector3D(1.0, 2.0, 4.0)

assert vector.coord_x == approx(1), f'{vector.coord_x=} != 1'
assert vector.coord_y == approx(2), f'{vector.coord_y=} != 2'
assert vector.coord_z == approx(4), f'{vector.coord_z=} != 4'
```

### b. instance methods

- length() : computes length of vector* 
- renormalize()   : renormalizes vector to unit length vector*

<!-- grade_id: cell-3328f70c25343894, points: 1, grade: true, solution: false -->
```python
# Verify length method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

assert vector.length()**2 == approx(21), f'{vector.length()**2=} != 21'
```

<!-- grade_id: cell-96e5a932f3347dcf, points: 1, grade: true, solution: false -->
```python
# Verify norm method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

vector.renormalize()

assert vector.length() == approx(1), f'{vector.length()=} != 1'
```

### c. Overloaded operations  

- addition operator(+)     : implement `__add__` method 
- substraction operator(-) : implement `__sub__` method 

    

<!-- grade_id: cell-1f019ebe8894fb2b, points: 1, grade: true, solution: false -->
```python
#Verify addition operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 3.3, 0.9)

cvector = avector + bvector

assert cvector.coord_x == approx(3.0), f'{cvector.coord_x=} != 3'
assert cvector.coord_y == approx(5.3), f'{cvector.coord_y=} != 5.3'
assert cvector.coord_z == approx(4.9), f'{cvector.coord_z=} != 4.9'

```

<!-- grade_id: cell-965fcb9cfdf9230b, points: 1, grade: true, solution: false -->
```python
#Verify substraction operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 0.2, 0.9)

cvector = avector - bvector

assert cvector.coord_x == approx(-1.0), f'{cvector.coord_x=} != -1'
assert cvector.coord_y == approx(1.8), f'{cvector.coord_y=} != 1.8'
assert cvector.coord_z == approx(3.1), f'{cvector.coord_z=} != 3,1'
```

## 2. Create class for a multi-currency bank account 


### 2.a

The account objects have an instance attribute that is initialized on creation

    >>> mca = MultiCurrencyAccount('SEK')
    >>> mca.base_currency
    'SEK'
    
### 2.b

We can model current exchange rates with a provided dictionary

    >>> exchange_rates['USDSEK']
    10
    
Use that dictionary to define a method in the class that converts a given currency to the base currency

    >>> mca.conversion_rate('USD')
    10
    
### 2.c

Define an instance attribute as a dictionary that holds different currencies that is initialized to zero for the base currency and a deposit method that updates it

    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0})
    >>> mca.deposit(10, 'USD')
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})

    
For this use a special version of dictonary from the collections module: defaultdict. If an element is not defined it is automatically created

### 2.d

Define a method `total_balance` that calculated the total value of the balances in the base currency

    >>> mca.deposit(10, 'USD')
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})
    >>> mca.total_balance()
    100
    
### 2.e

Consider a customer who needs to extract all his money to pay his electrical bill in the base currency.
A method `convert_to_base_currency` does the internal exchange before they can receive the mony


    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})
    >>> mca.convert_to_base_currency()
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 100, 'USD': 0})

<!-- grade_id: cell-cd199a6c4ed7bd19, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict

exchange_rates = {'USDSEK': 10, 'EURSEK': 10}

class MultiCurrencyAccount:
    """
    A multi-currency account with attributes:
    
        base_currency (str): international three-letter symbol
        balances (dict[str: int]): balance of each currency that has been introducted
        conversion_rate: instance method that returns exchange rate relative to base currency
        total_balance: value of all assets relative to base currency
        deposit: add funds in a given currency
        convert_to_base_currency: convert all funds to base currency
    """
    # YOUR CODE HERE
    ################
    def __init__(self,base_currency):
        self.base_currency = base_currency
        self.balances = defaultdict(int,{base_currency:0})
     
    def conversion_rate(self,currency):
        return exchange_rates.get(f'{currency}{self.base_currency}',1)

    def deposit(self,amount,currency):
        self.balances[currency] += amount
    
    def total_balance(self):
        total = sum(self.balances[currency]* self.conversion_rate(currency) for currency in self.balances )
        return round(total)

    def convert_to_base_currency(self):
        for currency, amount in self.balances.items():
            converted_amount = amount * self.conversion_rate(currency)
            self.balances[currency] = 0
            self.balances[self.base_currency] += converted_amount
            
```

<!-- grade_id: cell-cd199a6c4ed7bd19, points: 0, grade: false, solution: true -->
```python

```

```python
mca.balances
print(mca.balances)
# mca.deposit(100,'usd')
print(mca.balances)
```

Output:
```
defaultdict(<class 'int'>, {'SEK': 30, 'USD': 0, 'EUR': 0})
defaultdict(<class 'int'>, {'SEK': 30, 'USD': 0, 'EUR': 0})
```

### 2.a

The account objects have an instance attribute that is initialized on creation

    >>> mca = MultiCurrencyAccount('SEK')
    >>> mca.base_currency
    'SEK'

<!-- grade_id: cell-743e300b258a8a24, points: 1, grade: true, solution: false -->
```python
mca = MultiCurrencyAccount('SEK')
assert mca.base_currency == 'SEK', f'{mca.basecurrency=} != SEK'

```

### 2.b

We can model current exchange rates with a provided dictionary

    >>> exchange_rates['USDSEK']
    10
    
Use that dictionary to define a method in the class that converts a given currency to the base currency

    >>> mca.conversion_rate('USD')
    10


<!-- grade_id: cell-75ceff37da3ad5d5, points: 1, grade: true, solution: false -->
```python
mca = MultiCurrencyAccount('SEK')
assert mca.conversion_rate('USD') == 10, f"{mca.conversion_rate('USD')=} != 10"
```

### 2.c

Define an instance attribute as a dictionary that holds different currencies that is initialized to zero for the base currency and a deposit method that updates it

    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0})
    >>> mca.deposit(10, 'USD')
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})

    
For this use a special version of dictonary from the collections module: defaultdict. If an element is not defined it is automatically created



<!-- grade_id: cell-cfb2d9bd7b05f848, points: 1, grade: true, solution: false -->
```python
mca = MultiCurrencyAccount('SEK')
assert mca.balances == {'SEK': 0}, f'{mca.currencies=}'

mca.deposit(1, 'USD')
assert mca.balances == {'SEK': 0, 'USD': 1}, f'{mca.balances=}'

mca.deposit(2, 'EUR')
assert mca.balances == {'SEK': 0, 'USD': 1, 'EUR': 2}
```

### 2.d

Define a method `total_balance` that calculated the total value of the balances in the base currency

    >>> mca.deposit(10, 'USD')
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})
    >>> mca.total_balance()
    100
    

<!-- grade_id: cell-0564374e5e9da16a, points: 1, grade: true, solution: false -->
```python
mca = MultiCurrencyAccount('SEK')
assert mca.total_balance() == 0
mca.deposit(1, 'USD')
mca.deposit(2, 'EUR')

assert mca.total_balance() == 30, f'{mca.total_balance()=} != 30'
```

### 2.e

Consider a customer who needs to extract all his money to pay his electrical bill in the base currency.
A method `convert_to_base_currency` does the internal exchange before they can receive the mony


    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 0, 'USD': 10})
    >>> mca.convert_to_base_currency()
    >>> mca.balances
    defaultdict(<class 'int'>, {'SEK': 100, 'USD': 0})

<!-- grade_id: cell-ea6fce2baafa1212, points: 1, grade: true, solution: false -->
```python
mca.convert_to_base_currency()
assert mca.balances == {'SEK': 30, 'USD': 0, 'EUR': 0}
assert mca.total_balance() == 30, f'{mca.total_balance()=} != 30'
```

## C. Animals

Complete the Animal class to have and a class variable `animals` (list) and two instance variables, `name` (str) and `number` (int). You need to implement `__init__` and `__str__` methods

<!-- grade_id: cell-f9faff46e9639b35, points: 0, grade: false, solution: true -->
```python
class Animal:
    """
    A class for storing animals
    
    class attributes:
        animals: (list) to store all animals
    instance attributes:
        name:  (str) to store animal name
        number: (int) to store animal order number (starting with 1)
        
    class methods:
        __str__: string representation of animal, e.g. "1. Dog"
        
    static methods:
        zoo: returns string representation of all animals in orderd lies, e.g. 
           '''\
           1. Dog
           2. Cat'''
    """
    # YOUR CODE HERE
    ################
    animals = []

    def __init__(self,name):
       Animal.animals.append(self)
       self.name = name
       self.number = len(Animal.animals)
    
    def __str__(self):
         return f"{self.number}. {self.name.capitalize()}"

    def zoo():
       return '\n'.join(str(an) for an in Animal.animals)

```

<!-- grade_id: cell-41580406aa20538d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

dog = Animal('dog')
assert dog.name == 'dog', error_message(dog.name, 'dog')
assert dog.number == 1, error_message(dog.number, 1)
assert str(dog) == '1. Dog', error_message(str(dog), '1. Dog')

cat = Animal('cat')
assert cat.name == 'cat', error_message(cat.name, 'cat')
assert cat.number == 2, error_message(cat.number, 2)
assert str(cat) == '2. Cat', error_message(str(cat), '2. Cat')

```

A static method is a function in a class that is like an ordinary function, that does not depend on any instance (no self argument). 

In a class it is defined with a `@staticmethod` decorator.

It can be appended to the class definition as below. Complete the static method so that it returns a string which lists all member animals

<!-- grade_id: cell-a89d339b99865f5d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

#Generate a list of animals and compare with the class attribute
animals = [Animal(a) for a in ['camel', 'donkey', 'hippo']]
assert animals == Animal.animals, f'{animals} != {Animal.animals}'

#zoo should produce a printout of the defined animals
zoo_output = Animal.zoo()
print(zoo_output)
expected_output = """
1. Camel
2. Donkey
3. Hippo
"""

condition = zoo_output.strip() == expected_output.strip()
error_message = f"\n{zoo_output}\n\n   !=\n{expected_output}"

assert condition, error_message
```

Output:
```
1. Camel
2. Donkey
3. Hippo
```

```python
Animal.animals
```

Output:
```
[<__main__.Animal at 0x1055b85f0>,
 <__main__.Animal at 0x1055badb0>,
 <__main__.Animal at 0x1055b9430>]
```

```python

```