Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Yanjun Zhang"

```

---

# Advanced concepts


This lab provides exercises for decorators and iterators. Complete the assignments and return this notebook and return to Canvas

```python
import colorama
```

## Context managers

Context managers are statements with a `with` block and are implemented by classes with `__enter__` and `__exit__` methods. The official documentation is at https://docs.python.org/3/library/stdtypes.html#context-manager-types

### Assignment 1

Consider the color codes from the colorama package
They have the property illustrated by the following block:

```python
from colorama import Fore
GREEN = Fore.GREEN
RED = Fore.RED
RESET = Fore.RESET
print (f'{GREEN}yes {RED}no {RESET}maybe')
```

Output:
```
[32myes [31mno [39mmaybe
```

```python

```

Complete a context manager that changes the color of the output such that

<code>
>>> with Color(RED):
...     print('hello')    
<span style="color:red">hello</span>
>>> with Color(GREEN):
...     print('hello')    
<span style="color:green">hello</span>
>>> print('Back to normal')
Back to normal
</code>


<!-- grade_id: cell-e1388e8c145cd2ca, points: 0, grade: false, solution: true -->
```python
class Color:
    def __init__(self, color):
        # YOUR CODE HERE
        ################
        self.color = color
                
    def __enter__(self):
        # YOUR CODE HERE
        ################
        print(f'{self.color}')
                
    def __exit__(self, *args):
        print(RESET, end='')
```

```python
with Color(GREEN):
    print('Red day')
```

Output:
```
[32m
Red day
[39m
```

<!-- grade_id: cell-73773482ff268d10, points: 1, grade: true, solution: false -->
```python
# Verifies that you get color output


with Color(RED):
    print('Red day')
    
with Color(GREEN):
    print('Green forest')
    
print("Back to normal")
    
```

Output:
```
[31m
Red day
[39m[32m
Green forest
[39mBack to normal
```

## Decorators

A decorator is a function that modifies the behaviour of another function. The syntax

~~~
@decorator1
@decorator2
def f():
   ...
~~~

is so called syntactic sugar for

~~~
def f():
    ...
~~~

f = decorator1(decorator2(f))

The first form is convenient for our own code, but the second form must be used when we do not have the source at hand, e.g. for built-in functions

### Assignment 2

Repeat the functionality of Assignment1 for decorators such that a decorated function gets color output:

<code>
@print_in_red
def hello():
    print("Hello world!")
</code>

<code>
>>> hello()
<span style="color:red">Hello world!</span>
</code>
>>>

<!-- grade_id: cell-1282ae349daeef71, points: 0, grade: false, solution: true -->
```python
def print_in_red(f):
    # YOUR CODE HERE
    ################   
    def wrapper(*args, **kwargs):
        print(RED,'')  # Set color to red
        f(*args, **kwargs)
    return wrapper
    
```

```python
def print_in_red(f):
    # YOUR CODE HERE
    ################ 
 
    def wrapper(*args,**kwargs):
        
        print(RED,'')
        f(*args)
        
    return wrapper
```

<!-- grade_id: cell-19eb7c83720e263e, points: 1, grade: true, solution: false -->
```python
# Verify that you get color output

@print_in_red
def hello():
    print("Hello world!")
     
hello()

```

Output:
```
[31m 
Hello world!
```

```python

```

## Iterators/generators

### Assignment 3

The `datetime` module contains a number of helper function to handle dates and times. Here we will use the `date` and `timedelta` objects. 

* `date` objects refers to date information with year/month/day and more
* `timedelta` objects are differences between two dates (or date-times). They can be added to date objects to obtain new dates in some future.

Example of how they work is

```python
import datetime
today = datetime.date.today()
today
```

Output:
```
datetime.date(2023, 12, 27)
```

```python
week = datetime.timedelta(days=7)
today + week
```

Output:
```
datetime.date(2024, 1, 3)
```

Write a generator that returns the dates a number of weeks from now, in string format using `isoformat` method

<!-- grade_id: cell-03b5401cc183bf8e, points: 0, grade: false, solution: true -->
```python
def date_weeks_ahead(starting_date, number_of_weeks):
    # YOUR CODE HERE
    ################
    import datetime
    list = []
    date_list = []
    week = datetime.timedelta(days=7)
    staring_date = today
    
    for i in range(number_of_weeks):
       list = staring_date+datetime.timedelta(weeks = i+1)
       iso_list = list.isoformat()
       date_list.append(iso_list)
    return date_list
    
```

<!-- grade_id: cell-b42fc950abbaa239, points: 1, grade: true, solution: false -->
```python
today = datetime.date(2022, 3, 4)
actual = list(date_weeks_ahead(today, 4))
expected = ['2022-03-11', '2022-03-18', '2022-03-25', '2022-04-01']
assert actual == expected, f'{actual} != {expected}'
```

```python

```