Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Yanjun Zhang"

```

---

## Lab: basic Python syntax


#### Exercise 1
Write a function 'mymax' that accepts two numbers as parameters and that prints the larger number to the screen. For this elementary exercise, do not use the built-in function 'max()' - just revert to elementary expressions and statements.

<!-- grade_id: cell-ee8c56c7cc588b03, points: 0, grade: false, solution: true -->
```python
def mymax(x, y):
    """
    Returns the maximum of x and y
    """
    # YOUR CODE HERE
    ################
    
    if x > y:
        return(x)
    else:
        return(y)
  
```

<!-- grade_id: cell-304328f9c7d4c051, points: 1, grade: true, solution: false -->
```python
# Verifies the function mymax

calculated = mymax(2, 6)
expected = 6
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = mymax(5, 3)
expected = 5
assert calculated == expected, f'{calculated=} != {expected=}'
```

<!-- grade_id: cell-3de8d447b6562bf0, points: 1, grade: true, solution: false -->
```python
# Verifies that built-in max has not been used

import unittest.mock
with unittest.mock.patch('__main__.max') as mockmax:
    mymax(2, 6)
mockmax.assert_not_called()
```

#### Exercise 2

Write a function to sum the elements of a sequence with a for loop. You can assume that it consists of numeric values. Do not use the builtin sum function

<!-- grade_id: cell-5ac58bb5e671db91, points: 0, grade: false, solution: true -->
```python
def mysum(sequence):
    """
    Returns the sum of the elements in sequence
    """
    # YOUR CODE HERE
    ################
    
    result=0
    for element in sequence:
        result += element
    return result
    
```

```python

```

<!-- grade_id: cell-24ac435261756312, points: 1, grade: true, solution: false -->
```python
# Verifies mysum

calculated = mysum([1, 2])
expected = 3
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = mysum((1, 1/2, 1/4))
expected = 1.75
assert calculated == expected, f'{calculated=} != {expected=}'

```

```python
calculated = mysum([1,9,5,5])
expected=20
assert calculated == expected, f'{calculated=} {expected=}' 
```

<!-- grade_id: cell-f3e385c711f70755, points: 1, grade: true, solution: false -->
```python
# Verifies that built-in sum has not been used

import unittest.mock
with unittest.mock.patch('__main__.sum') as mocksum:
    mysum((2, 6))
mocksum.assert_not_called()
```

```python

```

Output:
```
red
yellow
```

#### Exercise 3

Variables of `str` type has a method `join` defined such that when a sequence (e.g. `list`, `tuple`) of strings is given as an argument, they are concatenated to a single string with the variable as a separator. E.g.

    >>> separator = '-'
    >>> separator.join(['father', 'in', 'law'])
    'father-in-law'

    >>> separator = '\n'
    >>> separator.join(['father', 'in', 'law'])
    'father\nin\nlaw'
    
    >>> print(separator.join(['father', 'in', 'law']))
    father
    in
    law


Use the `join` to generate a histogram with stars such that 

    histogram([3, 4, 1]) 
    
returns

    ***
    ****
    *

Hint: make use of the fact that a character multiplied by a number gives

    >>> 3 * '*'
    ***
    
and generate a list within the function that can be passed to the join function

<!-- grade_id: cell-631ae09a389a66d4, points: 0, grade: false, solution: true -->
```python
def histogram(sequence):
    """
    Returns histogram in the form of characters
    
    >>> histogram([3, 4, 1])
    ***
    ****
    *
    
    """
    # YOUR CODE HERE
    ################
    histogram_lines=[
        '*' * i for i in sequence 
    ]
    
    print (histogram_lines)
    result= '\n'.join(histogram_lines)
    return result
    

print(f'Example: the histogram of [1, 2, 3] is \n{histogram([1, 5, 2])}')
```

Output:
```
['*', '*****', '**']
Example: the histogram of [1, 2, 3] is 
*
*****
**
```

```python
sequence = [1,20]

for i in sequence:
    c = 'h' * i   
print(c)

```

Output:
```
hhhhhhhhhhhhhhhhhhhh
```

<!-- grade_id: cell-901ea19debe3ee4c, points: 1, grade: true, solution: false -->
```python
calculated = histogram([])
expected = ""
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = histogram([1, 2])
expected = "*\n**"
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = histogram([1, 2, 3])
expected = "*\n**\n***"
assert calculated == expected, f'{calculated=} != {expected=}'
```

#### Exercise 4

Compute how many times a character is present in a string. Use a dictionary to represent the string. 

<!-- grade_id: cell-94b1ecc182583f9b, points: 0, grade: false, solution: true -->
```python
def character_count(text):
    """
    Returns character count of text
    Input: text of type str
    Output: dict that maps a character to a number count
    """
    # YOUR CODE HERE
    ################
    
    count={}
    
    for i in text:
        if i in count:
            count[i] += 1
        else:
            count[i] = 1
    
    return count
    

print(f"Example: Character count of 'ABBA': {character_count('AqBfBA')} \nExpected {{'A': 2, 'B': 2}}")
```

Output:
```
Example: Character count of 'ABBA': {'A': 2, 'q': 1, 'B': 2, 'f': 1} 
Expected {'A': 2, 'B': 2}
```

<!-- grade_id: cell-2a787760db753909, points: 1, grade: true, solution: false -->
```python
# Verification

calculated = character_count("ABBA")
expected = {"A": 2, "B": 2}
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = character_count("CCGGAAGAGCTTACTTAG")
expected = {"A": 5, "C": 4, "G": 5, "T": 4}
assert calculated == expected, f'{calculated=} != {expected=}'
```

#### Exercise 5

A simple encryption scheme known as Caesarian roll, is obtained by shifting characters a fixed amount. In Python we can obtain the order of a character with the `ord` function

    >>> ord('A')
    65
    
and there reverse operation

    >>> chr(65)
    'A' 
    
Use this in a function that takes a string and returns the encrypted representation

<!-- grade_id: cell-95b86ffef93462ec, points: 0, grade: false, solution: true -->
```python
def caesar_shift(text, n):
    """
    Returns an encrypted string by shifting input characters by n
    """
    # YOUR CODE HERE
    ################
    
    fin_text=""
    
    for i in text:
        shift_i=chr(ord(i)+n)
        fin_text += shift_i
    
    
    return fin_text
    
```

```python
c = caesar_shift('pig',3)
print(c)
```

Output:
```
slj
```

<!-- grade_id: cell-4403018cd26de943, points: 1, grade: true, solution: false -->
```python
# Verification

calculated = caesar_shift('a', 0)
expected = 'a'
assert calculated == expected, f'{calculated=} != {expected=}'

calculated = caesar_shift('ABBA', 2)
expected = 'CDDC'
assert calculated == expected, f'{calculated=} != {expected=}'
```

```python

```